---
evaluation_interval: 12h

rule_files:
  - ../rules/cert_expiry.yaml

tests:
  - interval: 1d
    input_series:
      - series: obmondo_monitoring{certname="jaina.enableit", alert_id="monitor::domains::cert_expiry"}
        values: 1x1000
      - series: probe_ssl_earliest_cert_expiry{domain="obmondo.com",certname="jaina.enableit"}
        values: 1036800x100       # Tuesday, January 13, 1970 12:00:00 AM
      - series: threshold::monitor::domains::cert_expiry_days{domain="obmondo.com",certname="jaina.enableit"}
        values: 10x100            # Alert before 10 days
      - series: obmondo_monitoring{certname="jacen.enableit", alert_id="monitor::domains::cert_expiry"}
        values: 0x1000
      - series: probe_ssl_earliest_cert_expiry{domain="obmondo.com",certname="jacen.enableit"}
        values: 1036800x100       # Tuesday, January 13, 1970 12:00:00 AM
      - series: threshold::monitor::domains::cert_expiry_days{domain="obmondo.com",certname="jacen.enableit"}
        values: 10x100            # Alert before 10 days

    alert_rule_test:
      - alertname: monitor::domains::cert_expiry
        eval_time: 8d
        exp_alerts:
          - exp_labels:
              severity: critical
              certname: jaina.enableit
              alert_id: monitor::domains::cert_expiry
              domain: obmondo.com
            exp_annotations:
              summary: "SSL certificate will expire soon for domain **obmondo.com**"
              description: SSL Certificate which is on jaina.enableit expires in 4d 0h 0m 0s
      # certificate still has more then 11 days to expire, and we alert if its
      # less then 10 days (given by threshold)
      - alertname: monitor::domains::cert_expiry
        eval_time: 1d
        exp_alerts: []
