---
evaluation_interval: 1m

rule_files:
  - ../rules/dellhw.yaml

tests:
  - interval: 1m
    input_series:
      - series: obmondo_monitoring{certname="jacen.enableit", alert_id="monitor::system::service::dellhw::disk_healthy"}
        values: 1x1000
      - series: dell_hw_storage_pdisk_status{certname="jacen.enableit", disk="hdd1"}
        values: 1x1000

    alert_rule_test:
      - alertname: 'monitor::system::service::dellhw::disk_healthy'
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: 'critical'
              certname: 'jacen.enableit'
              disk: hdd1
              alert_id: 'monitor::system::service::dellhw::disk_healthy'
            exp_annotations:
              summary: 'Disk status for disk **hdd1** on **jacen.enableit** has failed'
              description: 'Disk status for disk **hdd1** on **jacen.enableit** has failed'
