---
evaluation_interval: 1m

rule_files:
  - ../rules/monitoring_stack.yaml

tests:
  - interval: 1m
    input_series:
      - series: watchdog_alerts_total{exported_instance="test.abc", instance="obmondomon.enableit", job="goopsmondo"}
        values: 50+1x1000

    alert_rule_test:
      - alertname: monitor::monitoring_stack::watchdog_missing
        eval_time: 0h
        exp_alerts: []

      - alertname: monitor::monitoring_stack::watchdog_missing
        eval_time: 4h
        exp_alerts: []

  - interval: 1m
    input_series:
      - series: watchdog_alerts_total{exported_instance="test.abc", instance="obmondomon.enableit", job="goopsmondo"}
        values: 50+0x1000

    alert_rule_test:
      - alertname: monitor::monitoring_stack::watchdog_missing
        eval_time: 4h
        exp_alerts:
          - exp_labels:
              severity: critical
              alert_id: monitor::monitoring_stack::watchdog_missing
              exported_instance: test.abc
              instance: obmondomon.enableit
              job: goopsmondo
              certname: test.abc
            exp_annotations:
              summary: "Prometheus Stack instance **test.abc** has stopped sending watchdog alerts"
              description: "Prometheus Stack instance **test.abc** has sent less than 1 watchdog alert in the past 35 minutes."
