---
evaluation_interval: 1m

rule_files:
  - ../rules/node_count.yaml

tests:
  - interval: 1m
    input_series:
      - series: days_missing_node_count{cluster_certname="k8s01.obmondo", instance="172.20.0.7:8080", job="obmondo", namespace="example"}
        values: 5+0x60
      - series: days_missing_node_count{cluster_certname="k8s02.obmondo", instance="172.20.0.7:8080", job="obmondo", namespace="example"}
        values: 11+0x60

    alert_rule_test:
      - alertname: ClusterFailedToUpdateNodeCountThisMonth
        eval_time: 30m
        exp_alerts:
          - exp_labels:
              severity: warning
              instance: '172.20.0.7:8080'
              cluster_certname: 'k8s01.obmondo'
              namespace: example
              job: obmondo
              alert_id: ClusterFailedToUpdateNodeCountThisMonth
            exp_annotations:
              description: 'The cluster **k8s01.obmondo** has failed to update its node count for ( **5** ) days this month. This can impact the process of calculating the average node count for the month.'
              summary: 'The cluster has not added their respective node counts to the DB for fewer than 10 days this month.'

      - alertname: ClusterFailedToUpdateNodeCountForOverTenDays
        eval_time: 30m
        exp_alerts:
          - exp_labels:
              severity: critical
              instance: '172.20.0.7:8080'
              cluster_certname: 'k8s02.obmondo'
              namespace: example
              job: obmondo
              alert_id: ClusterFailedToUpdateNodeCountForOverTenDays
            exp_annotations:
              description: 'The cluster **k8s02.obmondo** has failed to update its node count for over ten days this month ( **11** ). This will have an impact on the average node count calculation for the month.'
              summary: 'The cluster has not added their respective node counts to the DB for more than 10 days this month.'
