---
evaluation_interval: 1m

rule_files:
  - ../rules/opsmondo.yaml

tests:
  - interval: 1m
    input_series:
      - series: obmondo_monitoring{certname="test.abc", alert_id="monitor::service::opsmondo::alert_without_subscription"}
        values: 1x1000
      - series: opsmondo_alert_without_subscription_total{certname="test.abc"}
        values: 87+0x10

    alert_rule_test:
      - alertname: monitor::service::opsmondo::alert_without_subscription
        eval_time: 10m
        exp_alerts: []

  - interval: 1m
    input_series:
      - series: obmondo_monitoring{certname="test.abc", alert_id="monitor::service::opsmondo::alert_without_subscription"}
        values: 1x1000
      - series: opsmondo_alert_without_subscription_total{certname="test.abc"}
        values: 87+5x100

    alert_rule_test:
      - alertname: monitor::service::opsmondo::alert_without_subscription
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: warning
              certname: test.abc
              alert_id: monitor::service::opsmondo::alert_without_subscription
            exp_annotations:
              summary: "Server **test.abc** is sending alerts without subscription"
              description: |+2
                Server test.abc is sending alerts but has no current
                subscription.

                This can be due to the server getting auto-removed after failing
                Puppet check-in for too long.

  - interval: 30m
    input_series:
      - series: obmondo_monitoring{certname="test.abc", alert_id="monitor::service::opsmondo::too_few_watchdog_alerts"}
        values: 1x1000
      - series: opsmondo_watchdog_alerts_total{certname="test.abc"}
        values: 87+1x10

    alert_rule_test:
      - alertname: monitor::service::opsmondo::too_few_watchdog_alerts
        eval_time: 2h
        exp_alerts: []

  - interval: 5m
    input_series:
      - series: obmondo_monitoring{certname="test.abc", alert_id="monitor::service::opsmondo::too_few_watchdog_alerts"}
        values: 1x1000
      - series: opsmondo_watchdog_alerts_total{certname="test.abc"}
        values: 87+0x6 88+0x20 89+0x20 90+0x10 91+0x6

    alert_rule_test:
      - alertname: monitor::service::opsmondo::too_few_watchdog_alerts
        eval_time: 2h
        exp_alerts:
          - exp_labels:
              severity: warning
              certname: test.abc
              alert_id: monitor::service::opsmondo::too_few_watchdog_alerts
            exp_annotations:
              summary: "Server **test.abc** has stopped sending watchdog alerts"
              description: |+2
                Server test.abc have been sending less than 1 watchdog
                alert in the past 65 minutes (current average is 0/hour).

                It could mean that we're no longer accepting watchdog alerts or that
                the client is down.
