---
evaluation_interval: 1m

rule_files:
  - ../rules/prometheus.yaml

tests:
  # node down
  - interval: 1s
    input_series:
      - series: obmondo_monitoring{certname="example.jfk", alert_id="monitor::metrics::missing::pushprox"}
        values: 1+0x1800  # or 1x1800 in shorthand
      - series: up{certname="example.jfk", instance="example.jfk:63395", job="pushprox", pushprox_target="dns_exporter"}
        values: 0+0x1800
      - series: node_time_seconds{certname="example.jfk", instance="example.jfk:63395", job="pushprox", pushprox_target="dns_exporter"}
        values: 1668423358+1x1800  # system time should increase every time it's measured
      - series: node_boot_time_seconds{certname="example.jfk", instance="example.jfk:63395", job="pushprox", pushprox_target="dns_exporter"}
        values: 1661942053+0x1800

    alert_rule_test:
      - alertname: 'monitor::metrics::missing::pushprox'
        eval_time: 30m
        exp_alerts:
          - exp_labels:
              severity: critical
              instance: 'example.jfk:63395'
              certname: 'example.jfk'
              job: pushprox
              alert_id: monitor::metrics::missing::pushprox
              pushprox_target: dns_exporter
            exp_annotations:
              summary: |
                Missing Pushprox metrics for target dns_exporter on **example.jfk:63395**
              description: |
                Prometheus is missing Pushprox metrics for target dns_exporter from instance **example.jfk:63395**

  # node down
  - interval: 1m
    input_series:
      - series: obmondo_monitoring{certname="example.efg", alert_id="monitor::metrics::missing::pushprox"}
        values: 1x1000
      - series: up{certname="example.efg", instance="example.efg:63395", job="pushprox", pushprox_target="dns_exporter"}
        values: 1x30 0x30

    alert_rule_test:
      - alertname: 'monitor::metrics::missing::pushprox'
        eval_time: 1h
        exp_alerts: []

  # flapping
  - interval: 1m
    input_series:
      - series: obmondo_monitoring{certname="example.xyz", alert_id="monitor::metrics::missing::pushprox"}
        values: 1x1000
      - series: up{certname="example.xyz", instance="example.xyz:63395", job="pushprox", pushprox_target="dns_exporter"}
        values: 1x10 0x3 1x10 0x4 1x10

    alert_rule_test:
      - alertname: 'monitor::metrics::missing::pushprox'
        eval_time: 30m
        exp_alerts: []
